// variables1.rs
// 让我能够编译！执行 `rustex hint variables1` 获取提示 :)

// 关于 `I AM NOT DONE`：
// 我们有时会鼓励你在特定的练习中不断去尝试,
// 即使你已经得出了解决方案。如果你已经通过练习并做好了
// 应对下一题的准备, 请删除下面的 `I AM NOT DONE` 注释。

// I AM NOT DONE

/// 翻译: [mg-chao](https://github.com/mg-chao)
/// 加油 💪
fn main() {
    x = 5;
    println!("x has the value {}", x);
}
